Separating the control of moving and holding in post-stroke arm paresis

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

This study extends the previous interesting work of this group to address the potentially differential control of movement and posture. Their earlier work explored a broad range of data to make the case for a downstream neural integrator hypothesized to convert descending velocity movement commands into postural holding commands. Included in that data were observations from people with hemiparesis due to stroke. The current study uses similar data, but pushes into a different, but closely related direction, suggesting that these data may address the independence of these two fundamental components of motor control. I find the logic laid out in the second sentence of the abstract ("The paretic arm after stroke is notable for abnormalities both at rest and during movement, thus it provides an opportunity to address the relationships between control of reaching, stopping, and stabilizing") less then compelling, but the study does make some interesting observations. Foremost among them, is the relation between the resting force postural bias and the effect of force perturbations during the target hold periods, but not during movement. While this interesting observation is consistent with the central mechanism the authors suggest, it seems hard to me to rule out other mechanisms, including peripheral ones. These limitations should should be discussed.

Thank you for summarizing our work. Note we have improved the logic in our abstract (…”providing an opportunity to ask whether control of these behaviors is independently affected in stroke”) based on your comments as outlined in our previous revision. We now extensively discuss limitations and potential alternative mechanisms in greater detail, in a dedicated section (lines 846-895; see response to reviewer 2 for further details).

Reviewer #2 (Public review):

Summary:

Here the authors address the idea that postural and movement control are differentially impacted with stroke. Specifically, they examined whether resting postural forces influenced several metrics of sensorimotor control (e.g., initial reach angle, maximum lateral hand deviation following a perturbation, etc.) during movement or posture. The authors found that resting postural forces influenced control only following the posture perturbation for the paretic arm of stroke patients, but not during movement. They also found that resting postural forces were greater when the arm was unsupported, which correlated with abnormal synergies (as assessed by the Fugl-Meyer). The authors suggest that these findings can be explained by the idea that the neural circuitry associated with posture is relatively more impacted by stroke than the neural circuitry associated with movement. They also propose a conceptual model that differentially weights the reticulospinal tract (RST) and corticospinal tract (CST) to explain greater relative impairments with posture control relative to movement control, due to abnormal synergies, in those with stroke.

Thank you for the brief but comprehensive summary. We would like to clarify one point: we do not suggest that our findings are necessarily due to the neural circuitry associated with posture being more impacted than the neural circuitry associated with movement. (rather, our conceptual model suggests that increased outflow through the (ipsilateral) RST, involved in posture, compensates for CST damage, at the expense of posture abnormalities spilling over into movement). Instead, we suggest that the neural circuitry for posture vs. movement control remains relatively separate in stroke, with impairments in posture control not substantially explaining impairments in movement control.

Comments on revisions:

The authors should be commended for being very responsive to comments and providing several further requested analyses, which have improved the paper. However, there is still some outstanding issues that make it difficult to fully support the provided interpretation.

Thank you for appreciating our response to your earlier comments. We address the outstanding issues below.

The authors say within the response, "We would also like to stress that these perturbations were not designed so that responses are directly compared to each other ***(though of course there is an *indirect* comparison in the sense that we show influence of biases in one type of perturbation but not the other)***." They then state in the first paragraph of the discussion that "Remarkably, these resting postural force biases did not seem to have a detectable effect upon any component of active reaching but only emerged during the control of holding still after the movement ended. The results suggest a dissociation between the control of movement and posture." The main issue here is relying on indirect comparisons (i.e., significant in one situation but not the other), instead of relying on direct comparisons. Using well-known example, just because one group / condition might display a significant linear relationship (i.e., slope_1 > 0) and another group / condition does not (slope_2 = 0), does not necessarily mean that the two groups / conditions are statistically different from one another [see Figure 1 in Makin, T. R., & Orban de Xivry, J. J. (2019). Ten common statistical mistakes to watch out for when writing or reviewing a manuscript. eLife, 8, e48175.].

We agree and are well aware of the limitation posed by an indirect comparison – hence the language we used to comment on the data (“did not seem”, “suggest”, etc.). To address this limitation, we performed a more direct comparison of how the two types of perturbations (moving vs. holding) interact with resting biases. For this comparison, we calculated a Response Asymmetry Index (RAI):

Above, 𝑟𝐴 is the response on direction where resting bias is most-aligned with the perturbation, and 𝑟𝑂 is the response on direction where resting bias is most-opposed to the perturbation.

We calculated RAIs for two response metrics used for both moving and holding perturbations: maximum deviation and time to stabilization/settling time. For these two response metrics, positive RAIs indicate an asymmetry in line with an effect of resting bias.

The idea behind the RAI is that, while the magnitude of responses may well differ between the two types of perturbations, this will be accounted for by the ratio used to calculate the asymmetry. The same approach has been used to assess symmetry/laterality across a variety of different modalities, such as gait asymmetry (Robinson et al., 1987), the relative fMRI activity in the contralateral vs. ipsilateral sensorimotor cortex while performing a motor task (Cramer et al., 1997), or the relative strength of ipsilateral vs. contralateral responses to transcranial magnetic stimulation (McPherson et al., 2018). Notably, the normalization also addresses potential differences in overall stiffness between holding vs. moving perturbations, which would similarly affect aligned and opposing cases (see our response to your following point).

Figure 8 shows RAIs we obtained for holding (red) vs. moving/pulse (blue) perturbations. For the maximum deviation (left), there is more asymmetry for the holding case though the pvalue is marginal (p=0.088) likely due to the large variability in the pulse case (individual values shown in black dots). For time to stabilization/settling time (right) the difference is significant (p=0.0048). Together, these analyses indicate that resting biases interact substantially more with holding compared to movement control, in line with a relative independence between these two control modalities. We now include this panel as Figure 8, and describe it in Results (lines 587-611).

Note that even a direct comparison does not prove that resting biases and active movement control are perfectly independent. We now discuss these issues in more depth, in the new Limitations section suggested by the Reviewer (lines 836-849).

The authors have provided reasonable rationale of why they chose certain perturbation waveforms for different. Yet it still holds that these different waveforms would likely yield very different muscular responses making it difficult to interpret the results and this remains a limitation. From the paper it is unknown how these different perturbations would differentially influence a variety of classic neuromuscular responses, including short-range stiffness and stretch reflexes, which would be at play here.

Much of the results can be interpreted when one considers classic neuromuscular physiology. In Experiment 1, differences in resting postural bias in supported versus unsupported conditions can readily be explained since there is greater muscle activity in the unsupported condition that leads to greater muscle stiffness to resist mechanical perturbations (Rack, P. M., & Westbury, D. R. (1974). The short-range stiffness of active mammalian muscle and its effect on mechanical properties. The Journal of physiology, 240(2), 331-350.). Likewise muscle stiffness would scale with changes in muscle contraction with synergies. Importantly for experiment 2, muscle stiffness is reduced during movement (Rack and Westbury, 1974) which may explain why resting postural biases do not seem to be impacting movement. Likewise, muscle spindle activity is shown to scale with extrafusal muscle fiber activity and forces acting through the tendon (Blum, K. P., Campbell, K. S., Horslen, B. C., Nardelli, P., Housley, S. N., Cope, T. C., & Ting, L. H. (2020). Diverse and complex muscle spindle afferent firing properties emerge from multiscale muscle mechanics. eLife, 9, e55177.). The concern here is that the authors have not sufficiently considered muscle neurophysiology, how that might relate to their findings, and how that might impact their interpretation. Given the differences in perturbations and muscle states at different phases, the concern is that it is not possible to disentangle whether the results are due to classic neurophysiology, the hypothesis they propose, or both. Can the authors please comment.

It is possible that neuromuscular physiology may explain part of our results. However, this would not contradict our conceptual model.

Regarding Experiment 1, it is possible that stiffness would scale with changes in background muscle contraction as the reviewer suggests. Indeed, Bennett and al.(Bennett et al., 1992) used brief perturbations on the wrist to assess elbow stiffness, finding that, during movement, stiffness was increased in positions with a higher gravity load (and, in general, in positions where the net muscle torque was higher). However, during posture maintenance (like in our Experiment 1), they found that stiffness did not vary with (elbow) position or gravity load (two characteristics of our findings in Experiment 1):

“The observed stiffness variation was not simply due to passive tissue or other joint angle dependent properties, as stiffnesses measured during posture were position invariant. Note that the minimum stiffness found in posture was higher than the peak stiffness measured during movement, and did not change much with the gravity load.” (illustrated in Fig. 5 of that paper)

We thus find it very unlikely that stiffness explains the difference between the supported vs. unsupported conditions in Experiment 1.

Even if stiffness modulation between the supported vs. unsupported conditions could explain our finding of stronger posture biases in the latter case, it would not be incompatible with our interpretation of increased RST drive: increased stiffness would potentially magnify the effects of the RST drive we propose to drive these resting biases. It is possible that the increase in resting biases under conditions of increased muscle contraction (lack of arm support) is mediated through an increase in muscle stiffness. In other words, the increase in resting biases may not directly reflect additional RST outflow per se, but the scaling, through stiffness, of the same magnitude of RST outflow. Understanding this interaction was beyond the scope of our experiment design; in line with this, we briefly comment about it in our Limitations section.

Regarding Experiment 2, stiffness has indeed been shown to be lower during movement, and we now comment the potential effect of this on our results in the “Limitations” section (lines 815-830, replicated below). Importantly, for the case of holding perturbations, the increased stiffness associated with holding would increase resistance to both extension and flexion-inducing perturbations. Thus, higher stiffness would be unlikely to explain our finding whereby resting biases resist or aggravate the effects of holding perturbations depending on perturbation direction. In addition, the framework in Blum et al., that describes how interactions between alpha and gramma drive can explain muscle activity patterns, does not rule out central neural control of stiffness: “muscle spindles have a unique muscle-within-muscle design such that their firing depends critically on both peripheral and central factors” (emphasis ours). It may be, for example, that gamma motoneurons controlling muscle spindles and stiffness are modulated from input from the reticular formation, making this a mechanism in line with our conceptual model.

“Moreover, it has been shown that joint stiffness is reduced during movement compared to holding control (Rack and Westbury, 1974; Bennett et al., 1992). Along similar lines, muscle spindle activity – which may modulate stiffness – scales with extrafusal muscle fiber activity (such as muscle exertion involved in holding) and forces acting through the tendon (Blum et al., 2020). Such observations could, in principle, explain why we were unable to detect a relationship between resting biases and active movement control but we readily found a relationship between resting biases and active holding control: reduced joint stiffness during movement could scale down the influence of resting abnormalities. There are two issues with this explanation, however. First, it is debatable whether this should be considered an alternative explanation per se: stiffness modulation could be, in total or in part, the manifestation of a central movement/posture CST/RST mechanism similar to the one we propose in our conceptual model. For example, (Blum et al., 2020) argue that muscle spindle firing depends on both peripheral and central factors. Second, increased stiffness would not necessarily help detect differences in how active postural control responds to within-resting-posture vs. out-of-resting-posture perturbations. This is because an overall increase in stiffness would likely increase resistance to perturbations in any direction.”

The authors should provide a limitations paragraph. They should address 1) how they used different perturbation force profiles, 2) the muscles were in different states which would change neuromuscular responses between trial phase / condition, 3) discuss a lack of direct statistical comparisons that support their hypothesis, and 4) provide a couple of paragraphs on classic neurophysiology, such as muscle stiffness and stretch reflexes, and how these various factors could influence the findings (i.e., whether they can disentangle whether the reported results are due to classic neurophysiology, the hypothesis they propose, or both).

Thank you for your suggestion. We now discuss these points in a separate paragraph (lines 846895), bringing together our previous discussion on stretch reflexes, our description of different perturbation types, and the additional issues raised by the reviewer above.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

The authors have responded well to all my concerns, save two minor points.

Figure 2 appears to be unchanged, although they describe appropriate changes in the response letter.

Thank you for catching this error – we now include the updated figure (further updated to use the terms near/distant in place of proximal/distal).

I still take issue with the use of proximal and distal to describe the locations of targets. Taking definitions somewhat randomly from the internet, "The terms proximal and distal are used in structures that are considered to have a beginning and an end," and "Proximal and distal are anatomical terms used to describe the position of a body part in relation to another part or its origin." In any case, the hand does not become proximal just because you bring it to your chest. Why not simply stick to the common and clearly defined terms "near" and "distant"?

Point taken. We have updated the paper to use the terms near/distant.

Additional changes/corrections not outlined above

We now include a link to the data and code supporting our findings (https://osf.io/hufy8/). In addition, we made several minor edits throughout the text to improve readability, and corrected occasional mislabeling of CCW and CW pulse data. Note that this correction did not alter the (lack of) relationship between resting biases and responses to perturbations during active movement.

Response letter references

Bennett D, Hollerbach J, Xu Y, Hunter I (1992) Time-varying stiffness of human elbow joint during cyclic voluntary movement. Exp Brain Res 88:433–442.

Blum KP, Campbell KS, Horslen BC, Nardelli P, Housley SN, Cope TC, Ting LH (2020) Diverse and complex muscle spindle afferent firing properties emerge from multiscale muscle mechanics. Elife 9:e55177.

Cramer SC, Nelles G, Benson RR, Kaplan JD, Parker RA, Kwong KK, Kennedy DN, Finklestein SP, Rosen BR (1997) A functional MRI study of subjects recovered from hemiparetic stroke. Stroke 28:2518–2527.

McPherson JG, Chen A, Ellis MD, Yao J, Heckman C, Dewald JP (2018) Progressive recruitment of contralesional cortico-reticulospinal pathways drives motor impairment post stroke. J Physiol 596:1211–1225 Available at: https://doi.org/10.1113/JP274968.

Rack PM, Westbury D (1974) The short range stiffness of active mammalian muscle and its effect on mechanical properties. J Physiol 240:331–350.

Robinson R, Herzog W, Nigg BM (1987) Use of force platform variables to quantify the effects of chiropractic manipulation on gait symmetry. J Manipulative Physiol Ther 10:172–176.

Williams PE, Goldspink G (1973) The effect of immobilization on the longitudinal growth of striated muscle fibres. J Anat 116:45.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: Moving and holding still have been posited to be under separate control regimes for both eye and arm movements. The paretic arm after stroke exhibits different abnormalities during rest vs. movement, providing an opportunity to ask whether control of these behaviors is independently affected in stroke. Here, we quantified resting postural abnormalities in stroke patients by measuring their biases in force production as they held their hand still in various locations in a planar workspace, and then assessed the influence of these resting force biases on reaching in the same workspace. We found that patients had marked resting force biases at each location, even when the arm was supported. However, these biases did not transfer to arm-supported planar reaching movements: not during initial acceleration, not in response to mid-trajectory perturbations, and not during deceleration to a stop. Rather, the abnormal resting forces only appeared to switch on after a movement had fully stopped. These findings suggest that moving and holding are functionally separable modes of control. At the same time, we found that the resting biases mirrored characteristics of abnormal synergies during movement: they markedly decreased when arm support was provided; they were higher in more distal positions which require breaking out of flexion; and they scaled with the Fugl-Meyer score for the upper extremity (a measure of intrusion of abnormal synergies during active movement). These three shared features suggest a common mechanism for resting biases and abnormal synergies, which appears to be a contradiction given the functional separation of moving and holding observed in the same patients. To resolve this paradox, we propose a conceptual model that predicts a breakdown in the functional separation between reaching and holding when patients move in the absence of weight support. This conceptual model posits that synergies are the behavioral manifestation of a spillover of posture into movement. Mapping these functional systems onto anatomical and physiological details of lesioned substrate after stroke may provide implementation-level insight into how normal arm motor control is assembled.

Revised version: Moving and holding-still (holding) have been posited to be under separate control regimes for both eye and arm movements. The paretic arm after stroke exhibits different abnormalities during rest vs. movement, providing an opportunity to ask whether control of these behaviors is independently affected in stroke. Here, we quantified resting postural abnormalities in stroke patients by measuring their biases in force production as they held their hand still in various locations in a planar workspace, and then assessed the influence of these resting force biases on reaching in the same workspace. We found that patients had marked resting force biases at each location, even when the arm was supported. However, these biases did not transfer to arm-supported planar reaching movements: not during initial acceleration, not in response to mid-trajectory perturbations, and not during deceleration to a stop. Rather, the abnormal resting forces only appeared to switch on after a movement had fully stopped. These findings suggest that moving and holding are functionally separable modes of control. At the same time, we found that the resting biases mirrored characteristics of abnormal synergies during movement: they markedly decreased when arm support was provided; they were higher in more distant positions that require breaking out of flexion; and they scaled with the Fugl-Meyer score for the upper extremity (a measure of intrusion of abnormal synergies during active movement). These three shared features suggest a common mechanism for resting biases and abnormal synergies, which appears to be a contradiction given the functional separation of moving and holding observed in the same patients. To resolve this paradox, we propose a conceptual model that predicts a breakdown in the functional separation between reaching and holding when patients move in the absence of weight support. This conceptual model posits that synergies are the behavioral manifestation of a spillover of posture into movement. Mapping these functional systems onto anatomical and physiological details of lesioned substrate after stroke may provide implementation-level insight into how normal arm motor control is assembled.

Introduction

Earlier version: A longstanding idea in motor control is that moving and holding still (abbreviated to holding throughout this paper) rely upon separate control regimes. Such separation has been shown in the oculomotor system by the classic work of David Robinson (Robinson, 1970) who demonstrated that that while the controllers for moving and holding are neurally distinct, the hold controller integrates the motor commands that are generated by the move controller. It has been suggested that a separation between these two modes of control may extend to other effectors (Shadmehr, 2017). In particular, substantial behavioral (Ghez et al., 2007) and physiological (Kurtzer et al., 2005; Shalit et al., 2012) evidence supports the idea that a similar dissociation may govern the control of reaching and holding for the arm (Jayasinghe et al., 2022). Moreover, behavioral evidence suggests that commands that are generated during reaching are integrated to produce a component of the commands that are generated during subsequent holding (Albert et al., 2020). This has raised the possibility that like the control architecture of the oculomotor system, a holding controller may lie downstream of the reaching controller.

Revised version: A longstanding idea in motor control is that moving and holding-still (abbreviated to holding throughout this paper) rely upon separate control regimes. Such separation has been shown in the oculomotor system by the classic work of David Robinson (Robinson, 1970) who demonstrated that that while the controllers for moving and holding are neurally distinct, the hold controller integrates the motor commands that are generated by the move controller. It has been suggested that a separation between these two modes of control may extend to other effectors (Shadmehr, 2017). In particular, substantial behavioral (Ghez et al., 2007; Scheidt et al., 2011) and physiological (Kurtzer et al., 2005; Shalit et al., 2012) evidence supports the idea that a similar dissociation may govern the control of reaching and holding for the arm (Jayasinghe et al., 2022). Moreover, behavioral evidence suggests that commands that are generated during reaching are integrated to produce a component of the commands that are generated during subsequent holding (Albert et al., 2020). This has raised the possibility that like the control architecture of the oculomotor system, a holding controller may lie downstream of the reaching controller.

Earlier version: Patients with hemiparesis after stroke typically exhibit dissociable deficits in the control of reaching and holding within and across effectors (Garland et al., 2009; Levin, 1996; Mani et al., 2013; Schaefer et al., 2009; Trombly, 1992; Zackowski et al., 2004), making hemiparesis a potent model for understanding the interaction, or dissociation, between these two modes of motor control. The magnitude of reaching and holding abnormalities can be relatively large, making it easier to measure and compare them.

Revised version: Patients with hemiparesis after stroke typically exhibit dissociable deficits in the control of reaching and holding within and across effectors (Garland et al., 2009; Levin, 1996; Mani et al., 2013; Schaefer et al., 2009; Scheidt and Stoeckmann, 2007; Trombly, 1992; Zackowski et al., 2004), making hemiparesis a potent model for understanding the interaction, or dissociation, between these two modes of motor control. The magnitude of reaching and holding abnormalities can be relatively large, making it easier to measure and compare them.

Here, we focused on one particular aspect of holding: resting at a position. One of the most common and readily recognizable motor symptoms after stroke is abnormal resting posture (Twitchell, 1951). For example, the typical hemiparetic arm posture consists of flexion at the fingers, wrist, and elbow (Carr and Kenney, 1992) (Figure 1A). Prominent approaches for the treatment of hemiparesis have been based on the idea that abnormal resting posture has a direct deleterious effect upon movement control; they advocate for adjusting overall posture in ways that may minimize such effects (Bobath, 1982; Carr and Kenney, 1994, 1992).

Framework and experiment setup.

Earlier version: (A) A patient exhibiting a typical flexor posture at rest. Dashed arrows indicate elements of the posture: shoulder depression, arm adduction/internal rotation, elbow flexion. If one were to physically constrain the hand in a position away from the resting posture, the torques involved in each component of the abnormal resting posture translate to a force on the hand (sketched as a blue arrow); we thus designed an experiment to measure the resting force bias on the hand, as a marker of the overall postural abnormality. The goal was to compare resting postural force biases to active movement control in the same area (B). (C) Experiment setup. The participant holds the handle of the robotic arm; reach targets and cursor position are projected on a screen on top; for arm support, the participant’s arm is strapped on an armrest (c) connected to an air sled (a) which rests on the table. Air is provided through tube labeled (b). (D) Top-down view of setup, illustrating the different hand positions where resting postural forces were measured in Experiment 1 (open circles). Also shown are the five target positions used in the reaching and holding task for Experiment 2 (filled red circles). The gray box indicates the workspace depicted in Figure 2. (E) Examples of measurements of resting force biases. Each panel shows the evolution of resting forces during the 5s holding period for one participant (same participants as in Figure 2), taken at positions close to midline and distal from the body, under the same condition (paretic arm, arm support given). Solid lines indicate the force on the x-axis (positive values indicate forces towards the left), whereas dashed lines indicate the force on the y-axis (positive values indicate forces towards the body). The shaded area indicates the time window over which forces were averaged to estimate the resting bias, illustrating how resting biases were relatively stable by the 2s mark. Note that the third panel includes one trial (blue) which was rejected following visual inspection as described in Materials and Methods – Data Exclusion Criteria, due to instability in force production and movement during the hold period.

Revised version: (A) A patient exhibiting a typical flexor posture at rest. Dashed arrows indicate elements of the posture: shoulder depression, arm adduction/internal rotation, elbow flexion. If one were to physically constrain the hand in a position away from the resting posture, the torques involved in each component of the abnormal resting posture translate to a force on the hand (sketched as a blue arrow); we thus designed an experiment to measure the resting force bias on the hand, as a marker of the overall postural abnormality. The goal was to compare resting postural force biases to active movement control in the same area (B). (C) Experiment setup. The participant holds the handle of the robotic arm; reach targets and cursor position are projected on a screen on top; for arm support, the participant’s arm is strapped on an armrest (c) connected to an air sled (a) which rests on the table. Air is provided through tube labeled (b). (D) Top-down view of setup, illustrating the different hand positions where resting postural forces were measured in Experiment 1 (open circles). Also shown are the five target positions used in the reaching and holding task for Experiment 2 (filled red circles). The gray box indicates the workspace depicted in Figure 2. (E) Examples of measurements of resting force biases. Each panel shows the evolution of resting forces during the 5s holding period for one participant (same participants as in Figure 2), taken at positions close to midline and distant from the body, under the same condition (paretic arm, arm support given). Solid lines indicate the force on the x-axis (positive values indicate forces towards the left), whereas dashed lines indicate the force on the y-axis (positive values indicate forces towards the body). The shaded area indicates the time window over which forces were averaged to estimate the resting bias, illustrating how resting biases were relatively stable by the 2s mark. Note that the third panel includes one trial (blue) which was rejected following visual inspection as described in Materials and Methods – Data Exclusion Criteria, due to instability in force production and movement during the hold period.

Added in the revised version: Examples of resting postural force biases.

Added in the revised version: Shown are three stroke patients and one healthy control. Arrows indicate magnitude and direction of abnormal resting postural forces as measured at the hand at each location. Isoclines and the corresponding color levels provide a visual representation of where these biases would tend to direct the hand (from red towards blue). These isoclines represent different levels of the (spatial) integral of the posture bias-field (with zero referring to the isocline passing through the center position). Postural bias force vectors cross these isoclines perpendicularly. Please see Supplementary materials on how this visual aid was constructed. The red dots are the reach targets, with the center location circled (used in Experiment 2). Note how abnormalities in the paretic side are considerably stronger when arm support is removed. FM-UE: Fugl-Meyer score for the Upper Extremity (0-66).

Our main aim was to determine whether resting abnormalities bleed over into active movements in the post-stroke arm. This is of great interest because the ability to separate moving from holding may be precisely what is lost after stroke, with postural abnormalities contaminating voluntary movement. If this is indeed the case, it might lend credence to the idea in rehabilitation that treating resting abnormalities can benefit voluntary movement.

In Experiment 1, we assessed resting postural abnormalities by measuring resting postural force biases in patients with stroke using a planar workspace. We investigated how these force biases varied with arm position, presence of arm support, and overall motor impairment. In Experiment 2, we proceeded to assess patients’ control of reaching and holding in the same workspace. We separately investigated effects on initial reach and bringing the reach to a stop, as these two might be separately controlled (Ghez et al., 2007; Scheidt and Ghez, 2007). We also investigated active holding control after the movement was over, by examining responses to perturbations that attempted to move the arm off the target, in order to confirm that the same controller is engaged for both passive (as in Experiment 1) and active holding at the same position.

Materials and methods

Participants and Ethics Statement

Sixteen stroke patients (age: 58.5±17.8 [mean±standard deviation], nine female) and nine age-range matched healthy control participants (age: 62.6±15.2, six female) were recruited for this study. Table 1 shows details for each patient whereas Table 2 shows summary demographics and assessment metrics for patients and controls. Procedures were approved by the Johns Hopkins Institutional Review Board, and participants provided written informed consent in accordance with the Declaration of Helsinki.

Eligibility criteria

We recruited patients with hemiparesis due to stroke. To be eligible for the study, patients had to be adult, exhibit some movement with the affected arm, and be able to provide informed consent and understand the tasks involved. Exclusion criteria were marked cognitive impairment (assessed based on the Montreal Cognitive Assessment, MoCA, cutoff of 20); severe aphasia or ideomotor apraxia, neglect or hemianopia; and orthopedic or pain issues.

Task Details

Earlier version: Participants were seated in a chair and grasped the handle of a robotic arm; the handle moved either passively (by itself, Experiment 1) or actively (by the participant, Experiment 2) on the horizontal plane. In Experiment 2 and in two out of four conditions of Experiment 1, participants’ lower arm was supported using a custom-made air-sled (Figure 1C). Above the plane was a screen which blocked direct vision of the arm; on this opaque screen we continuously projected a cursor indicating hand position (diameter: 3 mm), as well as the currently active target (diameter: 10 mm). Handle position was recorded at 200Hz, whereas subject-produced forces on the handle were recorded using a 6-axis force transducer. Experiments typically began with the paretic arm (see specific details below); for healthy controls, an arm was randomly assigned as primary, with its schedule matching that of the paretic arm in stroke patients. We decided to prioritize the paretic arm, since the primary analyses in this paper involve comparing different metrics taken from the paretic arm. We note that this absence of paretic/non-paretic counterbalancing could potentially have a limiting effect on interpreting paretic vs. non-paretic comparisons; these comparisons, however, are secondary, and we see no reason why any ordering effects could have more than a minimal effect on them.

Revised version: Participants were seated in a chair and grasped the handle of a robotic arm; the handle moved either passively (by itself, Experiment 1) or actively (by the participant, Experiment 2) on the horizontal plane. In Experiment 2 and in two out of four conditions of Experiment 1, participants’ lower arm was supported using a custom-made air-sled (Figure 1C). Above the plane was a screen which blocked direct vision of the arm; on this opaque screen we continuously projected a cursor indicating hand position (diameter: 3 mm), as well as the currently active target (diameter: 10 mm). Handle position was recorded at 200Hz, whereas subject-produced forces on the handle were recorded using a 6-axis force transducer. Experiments typically began with the paretic arm (see specific details below); for healthy controls, an arm was randomly assigned as primary, with its schedule matching that of the paretic arm in stroke patients (this arm was used for comparisons with patients’ paretic data). We decided to prioritize the paretic arm, since the primary analyses in this paper involve comparing different metrics taken from the paretic arm. We note that this absence of paretic/non-paretic counterbalancing could potentially have a limiting effect on interpreting paretic vs. non-paretic comparisons; these comparisons, however, are secondary, and we see no reason why any ordering effects could have more than a minimal effect on them.

Experiment 1: Measuring resting postural abnormalities

Following a previous paradigm (Kanade-Mehta et al., 2023; Laczko et al., 2017; Simo et al., 2013), Experiment 1 assessed resting postural forces by having the robot passively move participants to different positions in a 2D workspace and hold them still in each position while it measured the forces they inadvertently exerted. Participants were instructed to maintain grasp on the robotic handle but otherwise rest and not resist the robot’s motion as it slowly (5s movement time) moved from one position to the next and held them still (for an additional 5s). The array of positions (see Figure 1) could vary from one participant to the next, with each position visited three times for each block. During the passive moving and holding, a 3 mm white cursor indicating handle position and a 10 mm yellow disk indicating the destination of the passive movement were displayed. Though not essential from the participant’s point of view, this allowed the experimenter to monitor the status of the experiment.

Each participant completed 4 blocks, two with each arm and in each arm support condition (air sled, no air sled); the typical order was {paretic, no support} -> {non-paretic, no support} - > {paretic, air sled support} -> {non-paretic, air sled support}; however in four individuals we completed the two paretic blocks first as their hand had to be secured to the handle (with self-adherent wrap) for a stable grasp.

Each block typically took less than 10 minutes to complete, with Experiment 1 lasting about 40 minutes including breaks.

Experiment 2: Assessing reaching control

Earlier version: Using the same workspace as Experiment 1, Experiment 2 assessed motor control in a reaching task. Participants made 10-cm point-to-point reaches across an array of 5 targets (diameter: 10mm) within the workspace located both distal and proximal to the body (Figure 3A), sampling 8 different movement directions. A white cursor (diameter: 3mm), indicating hand position, was visible throughout the experiment. Participants were instructed to try and stop at each target within a 600-800ms window after movement onset. At the end of movement, feedback was provided to indicate whether they were too fast (time <600ms, target turning red), too slow (time>800ms, target turning blue), or within the right time range (target “exploding” with a chirping sound).

Revised version: Using the same workspace as Experiment 1, Experiment 2 assessed motor control in a reaching task. Participants made 10-cm point-to-point reaches across an array of 5 targets (diameter: 10mm) within the workspace located both near and distant to the body (Figure 3A), sampling 8 different movement directions. A white cursor (diameter: 3mm), indicating hand position, was visible throughout the experiment. Participants were instructed to try and stop at each target within a 600-800ms window after movement onset. At the end of movement, feedback was provided to indicate whether they were too fast (time <600ms, target turning red), too slow (time>800ms, target turning blue), or within the right time range (target “exploding” with a chirping sound).

Added in the revised version: Average resting postural force biases.

Added in the revised version: (A) Average resting postural forces for the paretic (red) and non-paretic (cyan) arms of patients, as well as control participants (gray), illustrating how abnormal forces in the paretic arm are stronger in more distant targets and attenuated when arm support is provided (lighter shades). To average across left- and right-hemiparetic patients, left-arm forces were flipped left to right. (B) corresponding average resting postural force magnitudes and resting postural force directions (0 indicating the 3 o’clock direction, with negative values indicating clockwise directions). Errorbars indicate mean±SEM (circular mean±SEM in the case of movement directions).

The experiment was divided into blocks of 96 movements each (12 in each of the 8 movement directions). It began with three blocks with the paretic arm, followed by three more with the non-paretic arm, two more with the paretic arm, and ending with two blocks with the non-paretic arm. The first block with each arm was a familiarization block. Arm support (air sled) was provided throughout the experiment, and breaks were given between blocks.

Except for two participants, who performed Experiment 1 and Experiment 2 on different days due to limitations in their schedule, the entirety of each session – consisting of Experiment 1, Experiment 2 and standard assessments - took place on the same day and typically lasted about 3.5 hours with breaks given between the blocks as necessary.

Most trials (two-thirds) consisted of unperturbed movements to assess reaching control. In the remaining third of reaches, a 12N, 70 ms bell-shaped force pulse that was 12N at its peak was applied by the robot lateral to the ideal movement direction (i.e. the direction formed between the center of the start position and the center of the target) after participants reached 2cm away from the starting position (Fine and Thoroughman, 2006; Smith and Shadmehr, 2005). On half of these trials the perturbation was oriented leftwards with respect to the movement (counterclockwise pulse) and the other half rightwards (clockwise pulse).

A fraction of trials in each block (20/96) imposed a perturbation after movement in order to assess active holding control. For these trials, the holding time at the target was extended by 5 to 7 seconds, during which participants were instructed to hold still on the target (to remind them, the word “HOLD” were shown close to the target). During this extended hold period, a 6N force was gradually imposed over 2s in one of four different directions (45°, 135°, 225°, 315°), held constant for a pseudorandom time interval uniformly ranging from 3 to 5 seconds, and then abruptly released. We refer to this as “release perturbation” throughout the paper. Each block presented each position/release perturbation direction combination exactly once.

Thus, in summary, each 96-movement block consisted of 64 unperturbed movements and 32 movements perturbed with a force pulse (16 clockwise, and 16 counter-clockwise). For 20 out of the 96 movements in each block, the hold period was extended to test the hold perturbation (4 trials for each of the 5 target locations, each one of the 4 trials testing one perturbation direction as shown in Figure 7C).

A different perturbation was applied to movement (pulse) as compared to holding (release) so as to maximize disruption for each type of control, rather than to enable direct comparison between them. For holding control, instead of an abruptly-imposed force similar to the pulse used to probe movement control, we used a perturbation in which the arm first held against a 6N force which was then abruptly released. The rationale was that this force preloading would further engage active holding control and thus be better suited to examine it. In line with this, recent work that examined static hold under both types of perturbation – one in which a load is abruptly imposed vs. one in which a load is abruptly removed as we did here – found that the latter was more difficult and hence a better assay for unmasking a holding abnormality (Lowrey et al., 2019).

Data Analysis

Analysis was performed using MATLAB (Mathworks, Natick MA). For Experiment 1, we averaged resting postural forces for the last 3s of the 5s passive holding period for each trial. We excluded the first 2s in order to avoid measuring potential increases in muscle tone that may arise due to the velocity by which the robot passively brought the arm to each measurement position (i.e., avoiding potential effects of spasticity). Recent work using a similar task to examine resting force biases, found that the effects of the velocity at which the robot brought the arm to a position are present for 2s after the end of the passive movement, but then dissipate (Kanade-Mehta et al., 2023). To obtain a measure of the average resting postural force at each position for each individual and condition, we further averaged forces across the three visits to the same position. For comparisons, we focused on resting postural forces on the 5 positions shown in Figure 1D; these forces were obtained directly (when the exact positions were sampled for the individual) or through interpolation (7/16 patients and 3/9 controls).

For Experiment 2, movement onset was defined as the moment in which participants’ velocity from the starting position exceeded 3.5 cm/s, movement end was defined as the moment the participant was within the target and moving at a speed of less than 3.5 cm/s. Initial reaching angle was calculated between hand position at movement onset and 150ms later; endpoint reaching angle was calculated between the position of the hand when it crossed within 2cm of the target and 150ms later. In pulse trials, settling time was defined as the time taken from pulse onset to the first moment absolute lateral velocity dipped below 2cm/s and remained so for at least 100ms (or the movement ended). In release perturbation trials, settling time was defined as the time taken from perturbation onset (release of holding force) to the moment when velocity dipped below 2cm/s (and remained below that amount for at least 100ms) and the distance from the target was less than 2cm.

Data exclusion criteria

In Experiment 1, some trials (7.5%) were flagged as erroneous after visual inspection of force and movement profiles. Erroneous here refers to trials where forces appeared unstable and/or there was movement during the robot hold period (please see Figure 1E illustrating an example one such trial, blue curves on the third panel). We excluded these trials from our main analyses. To ensure that this exclusion did not bias our findings, we repeated our analyses including these trials and obtained similar results (Figures S5-E, S6-E, S7-E).

In Experiment 2, we excluded as outliers movements in which initial movement direction (150ms after movement onset) was ≥90° away from target direction. This excluded 0.95% of patients’ movements and 0.33% of controls’ movements.

Moreover, in some release perturbation trials, patients took a long time to reach the stabilization criterion described in the previous section; mistakenly, our setup limited its recording time to only the first 2s after force release. The exact time to stabilization thus could not be measured for these particular trials, so they had to be excluded from analysis. Though only 13.2±3.3% (mean±SEM) of paretic stabilization trials were thus excluded in the patient population (1.4±0.4% in their non-paretic side, 2.3±2.0% in controls), there were three patients for which excluded trials were 25% or more of all paretic trials. To ensure there are no systematic effects of this issue, we repeated the analysis of Figure 7F (a) by excluding these three patients altogether or (b) by assigning a value of 2.0 seconds to the affected trials. In both cases, we obtained results similar to our main analysis (Figure S7-5).

Stability of resting posture bias measurements in Experiment 1

Figure 1E shows a few examples of the within-trial evolution of force bias measurements during the 5s hold period, illustrating the high within-trial temporal stability in our dataset – particularly after the 2s mark, which begins the measurement window. To systematically assess the stability of our estimates of resting force biases across different trials, we estimated the variance attributed to different measurements at the same location (in contrast to the variance attributed to different locations, conditions, and patients). We found that this measurement variance was only 9.0% of the total variance for resting bias magnitude.

Statistical comparisons

Earlier version: In Experiment 1, we used an ANOVA to investigate any effect of conditions {Position (distal targets [furthest three positions] vs. proximal targets [closest two positions]), Support (with/without air sled), FM-UE (continuous)} and their interactions.

Revised version: In Experiment 1, we used an ANOVA to investigate any effect of conditions {Position (distant targets [furthest three positions] vs. near targets [closest two positions]), Support (with/without air sled), FM-UE (continuous)} and their interactions.

In Experiment 2, we used paired t-tests for the within-subject comparisons of outcome variables (For unperturbed reaching: initial reaching angle, endpoint reaching angle; for in-movement pulse perturbations: maximum deviation in the pulse direction, settling time; for holding release perturbations: maximum deviation in the perturbation direction; time to stabilization; path to stabilization) against the corresponding resting biases measured in Experiment 1. We performed two types of comparisons:

First, our main analyses took, for each individual, the two instances for which the resting biases would have the strongest opposing effects. For variables related to unperturbed reaching, we would take the movement direction for which the resting biases would be oriented the strongest CCW vs. the movement direction for which the resting biases would be oriented the strongest CW (two out of eight directions). The corresponding trials (32 for each of the two instances per participant) were averaged for each individual. For variables related to responses to pulse perturbation, we would take the two movement directions for which the resting bias at start would most oppose vs. most align with the pulse. The corresponding trials (8 for each of the two instances) were averaged for each individual. For variables related to responses to the release perturbation, we took the two perturbation direction/hold position combinations for which the resting bias on that position would most oppose vs. most align with the perturbation. The corresponding trials (4 for each of the two instances) were averaged for each individual.

Second, we used linear regression to calculate, for each individual, the sensitivities of outcome variables from Experiment 2 to the corresponding resting biases measured in Experiment 1 using data from all trials (Figures S5-2, S6-2, and S7-2), rather than the most extreme instances. We then tested whether the average of these per-individual sensitivities was significantly greater than zero using a paired t-test.

We used the circular statistics toolbox CircStat (Berens, 2009) to estimate the circular mean±SEM for the direction of resting biases shown in Figure 3B, right.

Fugl-Meyer assessments

Assessments were separately scored by AMH and KK with scores subsequently averaged (hence some scores having decimal values). For cases of substantial score differences (3 points or more) scores were again reviewed by both raters together.

Results

Patients exhibited resting force biases across the whole workspace

We first assessed post-stroke resting postural abnormalities across a 2D workspace. In Experiment 1, participants grasped the handle of a robotic arm, which passively moved their hand to a series of positions that sampled the workspace in front of them (the setup and example workspace are shown in Figure 1C,D). Participants were instructed to maintain grasp, but otherwise relax their arm and not resist the actions of the robotic arm. Once the participant’s arm was passively moved to a given position, the robotic arm would hold still for 5 seconds, enabling us to measure the forces involuntarily exerted by the participant as their hand was held at that location. Each experimental block consisted of three visits to 15-25 positions sampled for each participant, in random order. Participants completed four different experimental blocks: two with each arm, with or without arm weight support (provided by an air sled, Figure 1C).

Earlier version: Patients displayed abnormal postural force biases when the robot held the hand still at various locations across the workspace. Figure 2 shows examples of this for three patients and a typical healthy control participant. The non-paretic arm produced little to no postural force biases, whereas the paretic arm produced substantial postural force biases, particularly when the hand was held in a more distal position. The postural force biases were strongest when patients had to support the weight of the arm against gravity. Moreover, the patient with the highest degree of impairment (top left subplot of Figure 2, as assessed using the Fugl-Meyer score for the Upper Extremity, FM-UE (Fugl-Meyer et al., 1975)) exhibited the strongest resting postural abnormalities.

Revised version: Patients displayed abnormal postural force biases when the robot held the hand still at various locations across the workspace. Figure 2 shows examples of this for three patients and a typical healthy control participant. The non-paretic arm produced little to no postural force biases, whereas the paretic arm produced substantial postural force biases, particularly when the hand was held in a more distant position. The postural force biases were strongest when patients had to support the weight of the arm against gravity. Moreover, the patient with the highest degree of impairment (top left subplot of Figure 2, as assessed using the Fugl-Meyer score for the Upper Extremity, FM-UE (Fugl-Meyer et al., 1975)) exhibited the strongest resting postural abnormalities.

Removed in the revised version: Examples of resting postural force biases.

Removed in the revised version: Shown are three stroke patients and one healthy control. Arrows indicate magnitude and direction of abnormal resting postural forces as measured at the hand at each location. Isoclines and the corresponding color levels provide a visual representation of where these biases would tend to direct the hand (from red towards blue). These isoclines represent different levels of the (spatial) integral of the posture bias-field (with zero referring to the isocline passing through the center position). Postural bias force vectors cross these isoclines perpendicularly. Please see Supplementary materials on how this visual aid was constructed. The red dots are the reach targets, with the center location circled (used in Experiment 2). Note how abnormalities in the paretic side are considerably stronger when arm support is removed. FM-UE: Fugl-Meyer score for the Upper Extremity (0-66).

Earlier version: To aggregate results across all participants, we focused on abnormal resting postural force biases at five specific hand locations, illustrated by the red dots in Figure 2. These locations were chosen as they were within the sampled workspace of all participants (as this workspace could differ from one participant to the next due to different participants’ segment lengths), represented positions both distal and proximal to the body, and, importantly, they were also the movement targets used in Experiment 2. Figure 3A shows subject-averaged resting postural forces at each of these five positions and Figure 3B shows the corresponding force magnitudes and directions averaged across subjects and the five positions. There were two key observations: first, postural force biases were stronger at locations further away from the body (average magnitude – paretic side – without arm support: 6.5±1.3N for the three more distal targets vs. 3.5±0.5N for the two more proximal targets, t15=3.5, p=0.0030; with arm support: 3.3±0.5N for the three more distal targets vs. 1.9±0.3N for the two more proximal targets, t15=3.7, p=0.0022) and tended to point towards it (average direction – paretic side: −93.3±6.0° unsupported vs. −136.5±21.8° supported, with −90° directed towards the body); second, postural force biases were roughly halved in magnitude when arm support was provided (average magnitude – paretic side: 5.3±0.9N unsupported vs. 2.8±0.4N supported, t15=4.0, p = 0.0013).

Revised version: To aggregate results across all participants, we focused on abnormal resting postural force biases at five specific hand locations, illustrated by the red dots in Figure 2. These locations were chosen as they were within the sampled workspace of all participants (as this workspace could differ from one participant to the next due to different participants’ segment lengths), represented positions both near and distant to the body, and, importantly, they were also the movement targets used in Experiment 2. Figure 3A shows subject-averaged resting postural forces at each of these five positions and Figure 3B shows the corresponding force magnitudes and directions averaged across subjects and the five positions. There were two key observations: first, postural force biases were stronger at locations further away from the body (average magnitude – paretic side – without arm support: 6.5±1.3N for the three more distant targets vs. 3.5±0.5N for the two more near targets, t15=3.5, p=0.0030; with arm support: 3.3±0.5N for the three more distant targets vs. 1.9±0.3N for the two more near targets, t15=3.7, p=0.0022) and tended to point towards it (average direction – paretic side: -93.3±6.0° unsupported vs. -136.5±21.6° supported, with -90° directed towards the body); second, postural force biases were roughly halved in magnitude when arm support was provided (average magnitude – paretic side: 5.3±0.9N unsupported vs. 2.8±0.4N supported, t15=4.0, p = 0.0013).

Removed in the revised version: Average resting postural force biases.

Removed in the revised version: (A) Average resting postural forces for the paretic (red) and non-paretic (cyan) arms of patients, as well as control participants (gray), illustrating how abnormal forces in the paretic arm are stronger in more distal targets and attenuated when arm support is provided (lighter shades). To average across left- and right-hemiparetic patients, left-arm forces were flipped left to right. (B) corresponding average resting postural force magnitudes and resting postural force directions (0 indicating the 3 o’clock direction, with negative values indicating clockwise directions). Errorbars indicate mean±SEM (circular mean±SEM in the case of movement directions).

Examining the influence of resting postural force biases on active motor control

In Experiment 2, we sought to investigate whether resting postural force biases influence active reaching and holding in the same workspace. Specifically, we measured their effects upon initial reach direction (Figure 4, blue) and approach direction when bringing the reach to a stop (Figure 4, red), as it has been suggested that these two components are separately controlled (Ghez et al., 2007; Scheidt and Ghez, 2007); see also (Karst and Hasan, 1991; Lestienne, 1979). We also asked whether the same controller is engaged while the arm is passively (Experiment 1) vs. actively (Experiment 2) held at the same position. For this purpose, we examined active holding control after the movement was over (Figure 4, black), using perturbations that attempted to push the arm off the target.

Three aspects of active motor control that we tested in Experiment 2.

We separately examined the early part of the reaching movement (blue) and the late part, when the arm was coming to a stop (red). This was done by studying both unperturbed movements at different stages and movements that were perturbed with brief force pulses. In addition, we examined active holding control after the movement was over (black), using perturbations that tried to move the arm away from the held point. Shown is an example of trajectory and (absolute) velocity profiles from the reaching and coming- to-a-stop parts of a trial (left) and active holding against a perturbation after the trial was over (right).

Earlier version: Resting postural force biases did not affect the control of active reaching in the same workspace

Revised version: Resting postural force biases did not seem to affect the control of active reaching in the same workspace

We began by examining the control of active reaching. Participants made 10-cm point-to-point reaching movements within an array of five positions (filled circles in Figure 1D), for a total of 8 different movement directions (Figure 5A). Arm support was provided by the air sled (Figure 1C). Patients’ movements were generally impaired, taking more time and traveling a longer path to reach the target than controls (Time to target: 1.59±0.12s for patients’ paretic side vs. 0.82±0.03s for controls, p = 0.000063; Path traveled to target 13.4±0.6cm for patients’ paretic side vs. 10.8±0.2cm for controls, p = 0.0071, Fig. 5B,C).

Abnormal resting postural force biases do not interact with active reaching.

(A) Target array for Experiment 2 (movement task), illustrating the 5 start/end points of reaches and the 8 movement directions. (B) Example outwards trajectories (unperturbed trials) for a patient (cyan: non-paretic side; red: paretic side) and a healthy control (gray). (C) Subject-averaged reach performance based on either time (top) or path length to target (bottom) indicates impaired reaching control in patients’ paretic side. (D-F) Within-subject analysis of whether resting postural forces at movement start bias early movement towards their direction. (D) Sketch illustrating the concept behind this analysis. Assuming a movement from a start position subject to a strong rightwards resting bias (Fstart), will that translate to a corresponding rightwards movement bias which can be expressed as the early reach angle θstart,? (E) For each individual, we selected the target direction for which the counter-clockwise (CCW) component of Fstart was the strongest (red) vs. the target direction for which the clockwise (CW) component was the strongest (blue). The left panel shows this selection for an example participant: postural forces at start position were projected lateral to the movement direction, allowing us to select movement directions for which the lateral component was directed the strongest CCW or CW. The right panel shows the magnitude of these selected components across all patients. (F) Left: Corresponding movement trajectories (rotated so start position is at the bottom and end position at the top) for the directions selected for the same example participant. Right: Average initial angular deviations, θstart, across the selected directions for each participant. Note the lack of difference between the instances for which the CCW vs. the CW component of Fstart was the strongest, and thus no evidence that Fstart impinges upon the movement. (G-I) same as D-F but for endpoint resting postural forces, Fend and endpoint deviations, θend. Errorbars indicate SEM.

We then examined whether patients’ movements were affected by the resting postural biases measured in Experiment 1 in the same workspace and under the same arm-support condition. Would there, for example, be a difference when moving the arm through a high-postural bias area vs. a low-postural bias area? We examined effects of resting postural biases upon the initial reach vs. the approach to a hold position, as separate mechanisms may be involved in the control of each phase of movement (Ghez et al., 2007; Hannaford and Stark, 1985; Karst and Hasan, 1991; Sainburg et al., 1999; Scheidt and Ghez, 2007); resting postural biases might affect one phase but not the other. Specifically, we examined: (a) whether the direction of resting postural forces at the start position influenced trajectory deviations near the beginning of the movement (Figure 5D) or (b) whether direction of resting postural forces at the end position influenced trajectory deviations near the end of the movement (Figure 5G, after the participant reached within 2cm of the target).

To investigate within-subject effects, we selected, for each participant, the two movement directions for which the corresponding postural forces had strongest opposing effects – i.e. the most rightwards (CW) vs. most leftwards (CCW) with respect to movement direction. The idea behind picking the most extreme values for each individual was to maximize our sensitivity in detecting potential effects of resting postural forces upon active movement. Our analysis found no significant differences in directional biases between these two conditions: while the selected start-point force biases differed considerably between the most CCW and CW instances (1.9±0.4N vs. -2.1±0.4N, correspondingly [negative signs indicating CW forces], t15=5.68, p=0.000044), the corresponding initial angular deviations did not (0.4±2.1° for the most CCW vs. -1.6±2.2° for the most CW postural force, t15=0.56, p=0.58, Figure 5D-F). Similarly, while the selected end-point force biases also differed considerably (2.1±0.4N vs. - 1.9±0.4N, for the most CCW vs. CW instances, t15=5.68, p=0.000044), the corresponding endpoint angular deviations did not (3.6±2.3° for the most CCW vs. 4.9±2.6° for the most CW postural force, t15=0.50, p=0.62, Figure 5G-I). Figure S5-1 shows the same analyses but for non-paretic and control data, illustrating that, like in the paretic data, a lack of evidence that resting posture biases affect active movement.

While selecting the most extreme instances in terms of resting posture biases for each individual provides the greatest bandwidth to detect effects of such biases upon active movement, it also does not use the majority of the data (the 6 out of 8 movement directions corresponding to intermediate strengths of posture biases). We thus conducted an additional analysis which estimated the sensitivity (slope) of initial or endpoint deviations against the corresponding resting biases using the entirety of the data. In line with our main analysis, we found that patients’ slopes were not significant in either case (Initial reaching angle vs. posture bias at start: 1.5±1.4°/N, p = 0.30; endpoint reaching angle vs. posture bias at endpoint: - 1.0±0.8°/N, p = 0.25, Figure S5-2), further illustrating the lack of compelling evidence that postural abnormalities directly affect active movement.

Earlier version: A potential limitation in our data and analysis above is that unperturbed reaching movements may provide limited sensitivity in detecting effects of resting postural force biases, as any such effects may be largely compensated by a predetermined motor plan. To address this, we probed the effect of the biases to responses to less predictable, externally applied mid-movement force perturbations. In 1/3 of reaching movements in Experiment 2, chosen randomly, we imposed a 70ms bell-shaped force pulse which was 12N at its peak and acted lateral to the movement (Smith and Shadmehr, 2005). These pulses were imposed 2cm into the reach. Half of these pulses were clockwise (CW, blue in Figure 6A) and the other half were counter-clockwise (CCW, red in Figure 6A). We first verified that these pulses had a clear effect upon movement: perturbed movements took longer to complete in both patients (paretic movement time: 1.72±0.13s vs. 1.59±0.12s, t15 =5.31, p = 0.000088) and controls (movement time: 0.88±0.03s vs. 0.82±0.03s, t8 =5.56, p = 0.00053). Patients generally had impaired response to these pulses compared to their non-paretic side and healthy controls, deviating further (maximum lateral deviation – paretic: 2.5±0.1cm vs. non-paretic: 2.1±0.1cm [p= 0.0008] and controls: 2.1±0.1cm [p = 0.0244] - average of CW and CCW pulses) and taking a longer time to stabilize in the pulse direction (i.e. settling time; paretic: 0.61±0.03s vs. non-paretic: 0.47±0.01s [p=0.0027] and controls: 0.50±0.01s [p = 0.028]). This is illustrated in Figure 6B,C.

Revised version: A potential limitation in our data and analysis above is that unperturbed reaching movements may provide limited sensitivity in detecting effects of resting postural force biases, as any such effects may be largely compensated by a predetermined motor plan (Scheidt et al., 2011). To address this, we probed the effect of the biases to responses to less predictable, externally applied mid-movement force perturbations. In 1/3 of reaching movements in Experiment 2, chosen randomly, we imposed a 70ms bell-shaped force pulse which was 12N at its peak and acted lateral to the movement (Smith and Shadmehr, 2005). These pulses were imposed 2cm into the reach. Half of these pulses were clockwise (CW, blue in Figure 6A) and the other half were counter-clockwise (CCW, red in Figure 6A). We first verified that these pulses had a clear effect upon movement: perturbed movements took longer to complete in both patients (paretic movement time: 1.72±0.13s vs. 1.59±0.12s, t15 =5.89, p = 0.000030) and controls (movement time: 0.88±0.03s vs. 0.82±0.03s, t8 =5.56, p = 0.00054). Patients generally had impaired response to these pulses compared to their non-paretic side and healthy controls, deviating further (maximum lateral deviation – paretic: 1.25±0.05cm vs. non-paretic: 1.03±0.03cm [p= 0.00077] and controls: 1.07±0.04cm [p = 0.0357] - average of CW and CCW pulses) and taking a longer time to stabilize in the pulse direction (i.e. settling time; paretic: 0.61±0.03s vs. non-paretic: 0.48±0.01s [p=0.0027] and controls: 0.50±0.01s [p = 0.028]). This is illustrated in Figure 6B,C.

Responses to pulse perturbations during movement are not affected by resting postural force biases.

(A) Examples of perturbed (red: perturbed with CCW pulse; blue: perturbed with CW pulse) and unperturbed (gray) outward trajectories - same individuals as in Figure 3B. (B) Lateral velocity (positive: CCW to movement) before and after pulse onset, and corresponding responses from controls (gray), illustrating how patients, in response to the pulse, take longer time to settle and tend to experience larger lateral deviations compared to controls. (C) Summary performance measures for patients and controls, indicating impaired performance with the paretic side: settling time (left) and maximum lateral deviation on pulse direction (right). (D) Within-individual analysis: here, for each individual, we selected the movements for which the starting-position resting postural force would be either the strongest CCW or CW (left); we then examined the corresponding settling time (middle) and maximum lateral deviations (right). We find no effects of the most CCW vs. most CW resting postural forces in either case: there is no evidence for either reduced settling time or increased maximum lateral deviation for instances where pulse and resting bias are most opposing (open circles) compared to the instances where pulse and resting bias are most aligned (filled circles).

Added in the revised version: Responses to release perturbations during holding and their relationship to resting postural forces.

Added in the revised version: (A) Time course of the perturbation. (B) Example responses (all for the same position in the workspace) from two patients (top row) and two controls (bottom row). (C) Corresponding imposed force directions, the abrupt removal of which perturbs the movement in the opposite direction (compare with B). (D) Examples of tangential velocity profiles after the sudden release to the imposed hand force, averaged for all trials at the same position for each participant. Dashed line indicates the 2cm/s threshold used to assess time to stabilize. Left; example patient (paretic side); Right; example control. Colors correspond to different directions of the imposed hand force. (E) Summary of performance metrics after the perturbation for the paretic and non-paretic side of patients and healthy controls. (F) Within-subject analysis of the relationship between resting postural forces in the direction of the perturbation vs. performance against the perturbation. For each individual, we selected the two position/perturbation direction combinations for which resting postural forces were either the most opposed (green) to the perturbation or the most aligned (blue) with it. From left to right: forces in selected position/perturbation direction combinations; corresponding path traveled to stabilization; corresponding time to stabilization; corresponding maximum deviation. Note how the most-opposed resting bias for each patient is equal and opposite to the their most-aligned resting bias. This is because the same resting bias, when projected along the direction of two oppositely-directed perturbations (illustrated in C), it would oppose one with the same magnitude it would align with the other. This analysis suggests that resisting postural perturbations and restoring hand position after the perturbation was indeed easier when resting postural forces opposed, rather than were aligned with, the perturbation. Gray dots indicate individual data; colored dots and errorbars indicate mean±SEM.

Removed in the revised version: We then investigated whether resting postural forces played any role in patients’ response to the pulse perturbation. We reasoned that, should resting postural forces play a role, they would tend to decrease the effect of the pulse if they were in the opposite direction to it, and increase the effect of the pulse if they were in the same direction. This is because force pulses acting against the gradient of the postural bias field would act to stretch the already active muscles, further increasing postural resistance; whereas force pulses acting along the gradient of the postural bias field would act to shorten the same active muscles, reducing postural resistance. We thus compared the lateral component of resting postural forces in the start position against (a) the maximum lateral deviation in the direction of each pulse and (b) the time taken to stabilize in the pulse direction (settling time). In a manner similar to the previous section, we selected, for each individual, the directions where the starting lateral postural force was most

Earlier version: CCW vs. most CW, and compared the corresponding deviations. We found no clear effect: when resting postural forces were the most opposed to the pulse (vs. most aligned with it) there were no clear differences in deviation along the pulse direction, for neither the CCW pulses (0.97±0.17 vs. 1.30±0.14cm, t15 = 1.48, p = 0.16), the CW pulses (1.27±0.16 vs.1.50±0.20cm, t15 =0.87, p = 0.40), or with both types of pulses pooled together (1.12±0.12 vs. 1.40±0.11cm for the most-opposed vs. most-aligned pulse/bias instances, t15 =1.42, p = 0.18). Similarly, we found no clear differences in settling time between the instances where resting postural forces were the most opposed to the pulse vs. most aligned with it, for neither the CCW pulses (0.65±0.07 vs. 0.59±0.07s, t15 = 0.58, p = 0.57), the CW pulses (0.64±0.05 vs. 0.61±0.05s, t15 =0.40, p = 0.70), or with both types of pulses pooled together (0.64±0.05 vs. 0.60±0.05s, t15 =0.74, p = 0.47). For both metrics, we got similar findings when examining non-paretic and control data (Figure S6-1). In line with the above analyses, we also found no effects when we instead estimated the sensitivity of either outcome variable against the corresponding resting biases using data from all trials (for CCW and CW pulses combined, maximum lateral deviation along the pulse direction 0.07±0.06 cm/N, p=0.21; settling time: 0.00±0.02 s/N, p = 0.92, Figure S6-2).

Revised version: We then investigated whether resting postural forces played any role in patients’ response to the pulse perturbation. We reasoned that, should resting postural forces play a role, they would tend to decrease the effect of the pulse if they were in the opposite direction to it, and increase the effect of the pulse if they were in the same direction. This is because force pulses acting against the gradient of the postural bias field would act to stretch the already active muscles, further increasing postural resistance; whereas force pulses acting along the gradient of the postural bias field would act to shorten the same active muscles, reducing postural resistance. We thus compared the lateral component of resting postural forces in the start position against (a) the maximum lateral deviation in the direction of each pulse and (b) the time taken to stabilize in the pulse direction (settling time). In a manner similar to the previous section, we selected, for each individual, the directions where the starting lateral postural force was most CCW vs. most CW, and compared the corresponding deviations. We found no clear effect: when resting postural forces were the most opposed to the pulse (vs. most aligned with it) there were no clear differences in deviation along the pulse direction, for neither the CCW pulses (0.97±0.17 vs. 1.30±0.14cm, t15 = 1.48, p = 0.16), the CW pulses (1.27±0.16 vs.1.50±0.20cm, t15 =0.87, p = 0.40), or with both types of pulses pooled together (1.12±0.12 vs. 1.40±0.11cm for the most-opposed vs. most-aligned pulse/bias instances, t15 =1.42, p = 0.18). Similarly, we found no clear differences in settling time between the instances where resting postural forces were the most opposed to the pulse vs. most aligned with it, for neither the CCW pulses (0.61±0.05 vs. 0.64±0.05s, t15 = 0.40, p = 0.70), the CW pulses (0.59±0.07 vs. 0.65±0.07s, t15 =0.58, p = 0.57), or with both types of pulses pooled together (0.60±0.05s vs. 0.64±0.05, t15 =0.74, p = 0.47). For both metrics, we got similar findings when examining non-paretic and control data (Figure S6-1). In line with the above analyses, we also found no effects when we instead estimated the sensitivity of either outcome variable against the corresponding resting biases using data from all trials (for CCW and CW pulses combined, maximum lateral deviation along the pulse direction 0.07±0.06 cm/N, p=0.21; settling time: 0.00±0.02 s/N, p = 0.92, Figure S6-2).

In summary, we found no evidence for an interaction between resting postural force biases and the ability to respond to perturbations that were applied during movement.

Resting postural force biases emerged during active holding at the end of movement

Finally, in Experiment 2, we also investigated the relationship between resting postural force biases and active holding. In 20 out of 96 movements in each block, participants had to hold steady on the target for an additional 5 to 7 seconds (Figure 7A). During this time, the robot imposed a 6N force in one of four directions (45°, 135°, 225°, 315°, as shown in Figure 7B). This force was gradually applied over two seconds, held at a 6N level for 3-5 seconds, and then abruptly released, acting to displace the arm in the direction opposite to the original force, as illustrated in the examples in Figure 7B.

Removed in the revised version: Responses to release perturbations during holding and their relationship to resting postural forces.

Removed in the revised version: (A) Time course of the perturbation. (B) Example responses (all for the same position in the workspace) from two patients (top row) and two controls (bottom row). (C) Corresponding imposed force directions, the abrupt removal of which perturbs the movement in the opposite direction (compare with B). (D) Examples of tangential velocity profiles after the sudden release to the imposed hand force, averaged for all trials at the same position for each participant. Dashed line indicates the 2cm/s threshold used to assess time to stabilize. Left; example patient (paretic side); Right; example control. Colors correspond to different directions of the imposed hand force. (E) Summary of performance metrics after the perturbation for the paretic and non-paretic side of patients and healthy controls. (F) Within-subject analysis of the relationship between resting postural forces in the direction of the perturbation vs. performance against the perturbation. For each individual, we selected the two position/perturbation direction combinations for which resting postural forces were either the most opposed (green) to the perturbation or the most aligned (blue) with it. From left to right: forces in selected position/perturbation direction combinations; corresponding path traveled to stabilization; corresponding time to stabilization; corresponding maximum deviation. Note how the most-opposed resting bias for each patient is equal and opposite to the their most-aligned resting bias. This is because the same resting bias, when projected along the direction of two oppositely-directed perturbations (illustrated in C), it would oppose one with the same magnitude it would align with the other. This analysis suggests that resisting postural perturbations and restoring hand position after the perturbation was indeed easier when resting postural forces opposed, rather than were aligned with, the perturbation. Gray dots indicate individual data; colored dots and errorbars indicate mean±SEM.

Earlier version: Patients showed impaired capacity to resist and recover from this perturbation (the abrupt release of the imposed force). The time to stabilization for the paretic side (0.94±0.05s) was longer compared to the non-paretic side (0.79±0.03s, p = 0.024) and controls (0.78±0.06s, though this was statistically marginal, p = 0.061) as shown in Figure 7E, left. Moreover, patients traveled a longer path to stabilization (7.7±0.7cm for the paretic side vs. 5.9±0.3cm for the non-paretic side [p = 0.012] and 5.4±0.5cm for controls [p = 0.026], Figure 7E, middle) and they deviated more in the direction 2.6±0.1cm for the non-paretic side [p = 0.012] and 2.3±0.2 for controls [p = 0.0078], Figure 7E, right).

Revised version: Patients showed impaired capacity to resist and recover from this perturbation (the abrupt release of the imposed force). The time to stabilization for the paretic side (0.94±0.05s) was longer compared to the non-paretic side (0.79±0.03s, p = 0.024) and controls (0.78±0.06s, though this was statistically marginal, p = 0.061) as shown in Figure 7E, left. Moreover, patients traveled a longer path to stabilization (7.7±0.7cm for the paretic side vs. 5.9±0.3cm for the non-paretic side [p = 0.012] and 5.4±0.5cm for controls [p = 0.026], Figure 7E, middle) and they deviated more in the direction of the perturbation (3.3±0.2cm for the paretic side vs. 2.6±0.1cm for the non-paretic side [p = 0.012] and 2.3±0.2 for controls [p = 0.0078], Figure 7E, right).

To investigate whether resting postural force biases affected the control of actively holding across workspace locations, we performed a within-individual analysis analogous to the one we used for unperturbed reaches and moving perturbations. We first projected resting postural forces upon the directions of the release perturbation to assess the component of the resting postural force that opposed or aligned with the release perturbation. We then selected, for each participant, the two position/perturbation direction combinations for which these forces were either the most opposed (Figure 7F, green) to the perturbation or the most aligned (Figure 7F, blue) with it. For these selected position/perturbation direction combinations, we compared capacity to resist and recover from the perturbation, and found that this capacity was indeed better when the resting postural force was in a direction that opposed the perturbation (path traveled to stabilization: 3.7±1.0cm vs. 6.3±0.9cm, t15=2.8, p=0.014; time to stabilization: 0.6±0.1s vs. 0.9±0.1s, t15=3.8, p = 0.0017; maximum deviation: 1.7±0.4cm vs. 3.0±0.4cm, t15=4.6, p = 0.00036, corresponding to instances with the most opposed vs. the most aligned resting postural force). We did not observe significant differences when we performed the same analyses for non-paretic and control data (Figure S7-1).

We then proceeded, in a secondary analysis, to estimate the sensitivity of active holding control to resting postural biases for each individual using data from all trials (Figure S7-2). Specifically, we estimated the sensitivity of all three outcome variables (path to stabilization, time to stabilization, and maximum deviation) against the resistive component of the resting bias for each trial. Consistent with our main analysis in the previous paragraph, we found significant sensitivities in all cases (path to stabilization: sensitivity of -0.69±0.30 cm/N, p = 0.034; time to stabilization: sensitivity of -0.050±0.009 s/N, p = 0.000089; maximum deviation: -0.30±0.12 cm/N, p = 0.027). The negative slopes indicate reductions in these metrics – i.e. better performance – as the corresponding resting bias is more resistive. In short, our perturbations revealed that resting flexor biases “switched on” after movement was over, providing evidence for separate control between moving and holding.

Added in the revised version: Direct comparison of effects of resting postural force biases on moving vs. holding perturbations

Added in the revised version: The analysis above compared the effects of resting biases on moving vs. holding indirectly, finding effects in one case but not the other. We thus proceeded to directly compare how the two types of perturbations (moving vs. holding) interact with resting biases. An obstacle in this comparison is that the magnitude of perturbation responses may be different for moving vs. holding perturbations. This could be because the magnitude of the perturbations themselves is different, or because of other factors such as overall stiffness being different between the moving and holding state. To address this, we calculated a Response Asymmetry Index (RAI): Here, rA is the response in the direction where resting bias is most-aligned with the perturbation, and rO is the response in the direction where resting bias is most-opposed for the perturbation.

Added in the revised version: The idea behind the RAI is that, while the magnitude of responses may differ between the two types of perturbation, this will be accounted for by the ratio used to calculate the asymmetry. Similar approaches have been used to assess symmetry/laterality across a variety of different modalities (Cramer et al., 1997; McPherson et al., 2018; Robinson et al., 1987).

Added in the revised version: We calculated RAIs for two response metrics used for both types of perturbation: maximum deviation and time to stabilization/settling time. For the maximum deviation (Figure 8, left side), there is more asymmetry for the holding case though the p-value is marginal (p=0.088) likely due to the large variability in the pulse case (individual values shown in black dots). For time to stabilization/ settling time (Figure 8, right side) the difference is significant (p=0.0048). Together, these analyses indicate that resting biases interact substantially more with holding compared to moving, in line with a relative independence between these two control modalities.

Added in the revised version: Direct comparison of the effects of resting force biases on holding perturbations vs. pulse perturbations.

Added in the revised version: Response Asymmetry Indices (RAIs) are shown for holding (red) vs. moving/pulse (blue) perturbations. Positive values indicate a response asymmetry in line with an effect of resting force biases. Individual data shown in black dots. Errorbars indicate SEM.

Relating resting postural force biases to the Fugl-Meyer scale for the Upper Extremity

The observation, from Experiment 1, that resting postural force biases are reduced by external arm support bears parallels to the same effect of arm support upon abnormal synergies active during movement (Sukal et al., 2007). Here, “synergies” refer to abnormal co-activation patterns across joints that manifest as the patient tries to move – for example, the elbow involuntarily flexing as the patient tries to abduct their shoulder (Brunnstrom, 1966; Twitchell, 1951). Yet, Experiment 2 found no relationship between resting postural force biases and active movement control. To further investigate this apparent paradox, we examined the relationship between resting postural force biases and abnormal synergies in further detail. To assess the level of synergy abnormalities for each patient, we measured their Fugl-Meyer scores for the upper extremity (FM-UE), a scale which was designed to capture abnormal muscle synergy after stroke (Brunnstrom, 1966; Fugl-Meyer et al., 1975) and closely corresponds to EMG-based synergy measures (Bourbonnais et al., 1989).

Earlier version: We entered patients’ resting postural bias magnitudes into a mixed-effects ANOVA with FM-UE (continuous), Proximity (distal: the three locations furthest from the body, proximal: the two locations closest to the body, Figure 8, top left) and Support (with air sled, without air sled) as factors, evaluating for main effects and interactions. All factors showed significant main effects, with resting postural force magnitudes decreasing with FM-UE (R2 = 0.27, p 2×10-8), target proximity (η2 = 0.13, p = 0.00003), and weight support (η2 = 0.13, p = 0.00004). Significant interactions were observed between FM-UE and both Support (p=0.0062) and Proximity (p=0.0034), with Support and Proximity becoming more important for lower FM-UE scores (i.e. higher overall motor impairment / higher synergy intrusion) as illustrated in Figure 8.

Revised version: We entered patients’ resting postural bias magnitudes into a mixed-effects ANOVA with FM-UE (continuous), Proximity (distant: the three locations furthest from the body, near: the two locations closest to the body, Figure 9, top left) and Support (with air sled, without air sled) as factors, evaluating for main effects and interactions. All factors showed significant main effects, with resting postural force magnitudes decreasing with FM-UE (R2 = 0.27, p = 2×10-8), target proximity (η2 = 0.13, p = 0.00003), and weight support (η2 = 0.13, p = 0.00004). Significant interactions were observed between FM-UE and both Support (p=0.0062) and Proximity (p=0.0034), with Support and Proximity becoming more important for lower FM-UE scores (i.e. higher overall motor impairment / higher synergy intrusion) as illustrated in Figure 9.

Relationship between resting force biases and abnormal synergies.

Earlier version: Across-patient relationships of FM-UE (/66, higher scores indicating lower impairment) and resting postural force magnitudes, for distal (green) and proximal (blue) target positions, with (left) and without support (right). Note the strong effects of arm support, proximity, and FM-UE. Lines indicate linear fits; shading indicates 95% confidence interval for each fit.

Revised version: Across-patient relationships of FM-UE (/66, higher scores indicating lower impairment) and resting postural force magnitudes, for distant (green) and near (blue) target positions, with (left) and without support (right). Note the strong effects of arm support, proximity, and FM-UE. Lines indicate linear fits; shading indicates 95% confidence interval for each fit.

Earlier version: In summary, we made three key observations with regard to abnormal resting postural force biases. First, like abnormal synergies, they were exaggerated when active arm support was required. Second, they were more pronounced in more distal positions where the elbow was more extended; if patients were to actively reach to the same locations, they would have faced increased intrusion of flexor synergy. Third, they scaled with the synergy-based FM-UE. These observations suggest a common mechanism behind resting postural force biases and abnormal synergies.

Revised version: In summary, we made three key observations with regard to abnormal resting postural force biases. First, like abnormal synergies, they were exaggerated when active arm support was required. Second, they were more pronounced in more distant positions where the elbow was more extended; if patients were to actively reach to the same locations, they would have faced increased intrusion of flexor synergy. Third, they scaled with the synergy-based FM-UE. These observations suggest a common mechanism behind resting postural force biases and abnormal synergies.

Discussion

Earlier version: We assessed abnormal resting posture in stroke patients by measuring the resting force biases they involuntarily exerted while their arm was held at different points within a planar workspace (Experiment 1). We found that these resting postural force biases were strongest in more distal positions of the arm, generally pulled the arm toward a flexed position, and were significantly reduced, but nevertheless remained present, with arm weight support. We then proceeded to assess reaching and holding control in the same workspace in the presence of arm support (Experiment 2) and examined whether resting postural forces could partially account for deficits in the motor control of reaching and holding in the same patients. Remarkably, resting postural force biases did not have a detectable effect upon any component of active reaching, and only emerged during the control of holding after the movement ended. The results suggest a dissociation between the control of movement and posture. At the same time, assessing patients’ impairment using the FM-UE, a metric designed to measure for abnormal synergies during 3D arm movements, revealed a strong association between resting force biases and abnormal movement synergies; which raises the possibility that the observed dissociation of movement and posture control for planar weight-supported movements may break down for unsupported 3D arm movements. This dissociation raises interesting questions about both the neural architecture responsible for these two forms of control, and how to approach rehabilitation of the post-stroke arm.

Revised version: We assessed abnormal resting posture in stroke patients by measuring the resting force biases they involuntarily exerted while their arm was held at different points within a planar workspace (Experiment 1). We found that these resting postural force biases were strongest in more distant positions of the arm, generally pulled the arm toward a flexed position, and were significantly reduced, but nevertheless remained present, with arm weight support. We then proceeded to assess reaching and holding control in the same workspace in the presence of arm support (Experiment 2) and examined whether resting postural forces could partially account for deficits in the motor control of reaching and holding in the same patients. Remarkably, resting postural force biases did not have a detectable effect upon the control of active reaching, and only emerged during the control of holding after the movement ended. This suggests a dissociation between the control of movement and posture. At the same time, assessing patients’ impairment using the FM-UE, a metric designed to measure for abnormal synergies during 3D arm movements, revealed a strong association between resting force biases and abnormal movement synergies; which raises the possibility that the observed dissociation of movement and posture control for planar weight-supported movements may break down for unsupported 3D arm movements. This dissociation raises interesting questions about both the neural architecture responsible for these two forms of control, and how to approach rehabilitation of the post-stroke arm.

Dissociation between reaching and holding

Previous research provides evidence for separate control of reaching and holding in the healthy arm – for a review, see (Shadmehr, 2017) and (Jayasinghe et al., 2022). For example, following visuomotor rotation training specific to the outward phase of an out-and-back movement, participants did not transfer this rotation to holding after a point-to-point reach in the same direction: after the movement was over, and visual feedback was removed, participants’ held position drifted from the rotated movement endpoint towards the baseline hold position (Scheidt and Ghez, 2007). There is also neurophysiological evidence for separate control of moving and holding. A previous study found populations of neurons in macaque M1 that represent mechanical loads during posture or movement but not both (Kurtzer et al., 2005), whereas in another study cortical neurons and spinal interneurons coded parameters related to either movement or posture maintenance (Shalit et al., 2012).

The dissociation reported here between the control of reaching and holding in the post-stroke arm is consistent with a recently proposed hybrid model, which posits distinct controllers for reaching and for bringing the arm to a stop (Jayasinghe et al., 2022). Our finding that resting posture control does not interact either with the initial reach or bringing the arm to a stop extends this idea, suggesting three distinct controllers: one for the initial reach; another for bringing the arm to a stop; and another for control of holding at the endpoint after movement is over. It should be noted, however, that having distinct neural circuits for reaching and holding does not rule out interactions between them. For example, we recently demonstrated how arm holding control reflects the integration of motor commands driving the preceding active movement that led to the hold position, in both healthy participants and patients with hemiparesis (Albert et al., 2020). However, in that paper, we did not claim that this integration is the only source of holding control. Indeed, in Experiment 1 of the current study, we used passive movement to bring the arm to each probed position, which means that the postural biases could not be the result of integration of motor commands.

Potential origins of abnormal resting flexor biases

Experiment 1 revealed several parallels between resting postural force biases and abnormal synergies: a propensity for flexion, mitigation by arm weight support and, a significant correlation between patients’ resting postural force biases and their FM-UE scores. The FM-UE is a measure designed to assess post-stroke abnormal synergies during active movement. These similarities raise the possibility that post-stroke resting postural biases and movement synergies share a generative mechanism. Consistent with this possibility, a study found that externally-imposed elbow flexion led to (involuntary) shoulder flexion and external adduction in stroke patients but not healthy controls, suggesting abnormal synergy patterns do not require active voluntary movement to be expressed (Sangani et al., 2007) – also see (Schmit and Rymer, 2001).

It has been proposed that abnormal resting posture after brain injury can be ascribed to an extrapyramidal system (i.e., other than the corticospinal tract, CST) (Denny-Brown, 1964). Consistent with this suggestion, multiple lines of evidence from animal models point towards reticulospinal tract (RST) involvement in postural control (Deliagina et al., 2007; Lacquaniti et al., 1997; Takakusaki, 2017). Lesions of the monkey ponto-medullary reticular formation resulted in abnormal postures; notably, these were characterized by trunk and limb flexion, elevated shoulders, and arms held close to the body (Lawrence and Kuypers, 1968) mirroring the flexor posture pattern seen after stroke. Several studies in cats also suggest a postural role for the reticulospinal tract. Reticular formation neurons respond to vestibular inputs such as head tilts or whole-body tilts, in line with a role in adjusting posture against gravity (Bolton et al., 1992; Matsuyama and Drew, 2000; Pompeiano et al., 1984). Electrical and chemical stimulation of the reticular formation leads to modulations in muscle tone (Takakusaki et al., 2016), and reticulospinal neurons display tonic activity patterns that are related to postural adjustments that precede a movement rather than to the movement itself (Schepens and Drew, 2004).

There is also recent evidence that the RST is the descending system responsible for the generation of large forces during voluntary movement (Glover and Baker, 2022; Tapia et al., 2022). Thus, the accumulated evidence suggests that the RST could control posture and large force production in the upper limb. Upregulation of the RST has also been implicated in the generation of abnormal movement synergies after stroke (McPherson et al., 2018). For example, the ipsilateral RST can facilitate flexors but suppress extensors, a pattern mirroring the flexor synergy (Davidson et al., 2007; Davidson and Buford, 2004).

There appears to be a contradiction, however, between the idea that abnormal postures share a common mechanism with abnormal movement synergies, and our finding that there was no evidence for intrusion of resting biases into any phase of active reaching. A potential explanation for this apparent contradiction is that our experiments were conducted on a 2D surface with weight support. In contrast, the FM-UE scale, a synergy-based measure, is performed by patients in 3D without weight support. It has been shown that for 3D reaching, patients with chronic stroke express intrusive flexor synergies (Zackowski et al., 2004). Weight support – as was used in Experiment 2 – reduces intrusion of flexor synergies for planar movements (Sukal et al., 2007) and improves planar kinematics (Beer et al., 2007, 2004). Earlier work, using a weight-supported planar task, found similar degrees of reaching abnormality whether movements were made in or out of synergy (Levin, 1996); more recently, we showed that reaching dexterity can be dissociated from synergy intrusion when arm support is provided (Hadjiosif et al., 2022). Thus, while full weight support reduces both resting flexor biases and movement-related flexor synergies, this reduction is more complete for synergies compared to resting biases. This is not inconsistent with positing a shared substrate for abnormalities at rest and during movement, but this substrate might play a greater role in holding than in moving, and so the consequences of damage to it may be more apparent in a task that stresses holding over movement.

Earlier version: What would this framework look like? We posit that the motor system has separable functional modes for moving vs. holding and that this is accomplished by differentially weighting the contributions of descending systems that are operative in both modes (Figure 9). The CST is weighted more towards fast and fractionated control during movement, whereas the RST, in contrast, is weighted more towards slower postural control and generation of large isometric forces (such as vertical forces for arm support, or horizontal forces for holding the arm still against a background load like in our posture/release perturbation trials). In healthy individuals, these two modes are kept in balance by the CST, which has a moderating influence on the RST (Schepens and Drew, 2006) (Figure 9, left). In this framework, the CST is the controller during movement and the modulator during holding.

Revised version: What would this framework look like? We posit that the motor system has separable functional modes for moving vs. holding and that this is accomplished by differentially weighting the contributions of descending systems that are operative in both modes (Figure 10). The CST is weighted more towards fast and fractionated control during movement, whereas the RST, in contrast, is weighted more towards slower postural control and generation of large isometric forces (such as vertical forces for arm support, or horizontal forces for holding the arm still against a background load like in our posture/release perturbation trials). In healthy individuals, these two modes are kept in balance by the CST, which has a moderating influence on the RST (Schepens and Drew, 2006) (Figure 10, left). In this framework, the CST is the controller during movement and the modulator during holding.

An architecture for the separable control of reaching and holding and spillover effects in stroke.

Removed in the revised version: Post-stroke damage to the CST reduces the moderating influence the CST has upon the RST (Figure 9, right); influence that is likely further compromised by upregulation of the RST through plasticity mechanisms (García-Alías et al., 2015; Zaaimi et al., 2012). This conceptual model can explain our results as follows: As the CST is the dominant system during movement it can still modulate the RST in this mode, especially with weight support, as this reduces RST drive, with the consequence that resting biases do not markedly contaminate active movement. Conversely, the RST is the dominant system for postural control, and can overcome weakened

Earlier version: CST modulation, which leads to the resting biases we observed. Thus, weight support allows a weakened CST to keep moving protected from holding, but it cannot prevent abnormal holding itself. The interesting implication of this conceptual model is that synergies are in fact postural abnormalities that spill over into active movement when the CST can no longer modulate the increased RST activation that occurs when weight support is removed (i.e. resting biases may influence active reaching in absence of weight support). Supporting this idea, a study found increased ipsilateral activity (which primarily represents activation via the descending ipsilateral RST (Zaaimi et al., 2012)) when the paretic arm had reduced support compared to full support (McPherson et al., 2018).

Revised version: Post-stroke damage to the CST reduces the moderating influence the CST has upon the RST (Figure 10, right); influence that is likely further compromised by upregulation of the RST through plasticity mechanisms (García-Alías et al., 2015; Zaaimi et al., 2012). This conceptual model can explain our results as follows: As the CST is the dominant system during movement it can still modulate the RST in this mode, especially with weight support, as this reduces RST drive, with the consequence that resting biases do not markedly contaminate active movement. Conversely, the RST is the dominant system for postural control, and can overcome weakened CST modulation, which leads to the resting biases we observed. Thus, weight support allows a weakened CST to keep moving protected from holding, but it cannot prevent abnormal holding itself. The interesting implication of this conceptual model is that synergies are in fact postural abnormalities that spill over into active movement when the CST can no longer modulate the increased RST activation that occurs when weight support is removed (i.e. resting biases may influence active reaching in absence of weight support). Supporting this idea, a study found increased ipsilateral activity (which primarily represents activation via the descending ipsilateral RST (Zaaimi et al., 2012)) when the paretic arm had reduced support compared to full support (McPherson et al., 2018).

A side question is how the control of decelerating to a stop – which we show here is distinct from the control of holding after the movement is over – fits within such a scheme. A recent saccade study found that deceleration may be controlled through the cerebellum (Sedaghat-Nejad et al., 2022). In reaching tasks in the mouse, activity in the interpositus nucleus scaled with limb deceleration (Becker and Person, 2019) whereas disruption of the pontine nuclei did not impair movement initiation as much as kinematic variables related to bringing the movement to a stop (Guo et al., 2021).

Increased recruitment of the RST has also been implicated in power grip (Baker and Perez, 2017; Tazoe and Perez, 2017). This may be relevant to our study, as our task had participants actively grasp the handle of the robot for both Experiments 1 and 2. If increased RST recruitment indeed explains abnormal resting postural biases as we discussed in the previous section, a stronger grasp would in turn increase the strength of these biases even further. Moreover, the intermixing of free-reaching and perturbation trials in Experiment 2 could have led to increased uncertainty in environmental dynamics; uncertainty can lead to adjustments such as even stronger grip (Hadjiosif and Smith, 2015). Yet, despite this potential additional RST recruitment, we found that resting biases did not affect active reaching and only switched on after the reach was over; detected when we a applied a perturbing hand force. This result lends further support for separate controllers for reaching vs. holding.

Added in the revised version: Limitations

Another potential source for the resting postural force biases we observe could be abnormally low thresholds (Levin and Feldman, 1994; Powers et al., 1989) or abnormally high gains (Thilmann et al., 1991) of the stretch reflex. Passively extending the elbow, even at low speeds – something Experiment 1 did – can lead to increased muscle activity which may persist after the end of movement (Kanade-Mehta et al., 2023; Levin et al., 2000). This velocity-dependent increase in muscle tone – spasticity – could potentially explain some of the resting force biases measured in Experiment 1. However, in an earlier task where the elbow was passively extended at different velocities, the resulting tonic biceps EMG terminated around the time that the applied elbow extension was completed (Thilmann et al., 1991). And, in recent work, (Kanade-Mehta et al., 2023) used a paradigm similar to Experiment 1 and found that, while resting force biases may initially reflect the velocity at which the arm (passively) approached each test position, this velocity dependence dissipated 2s after the end of passive movement. In our study, resting biases were measured beginning 2s after the end of movement, which makes a central rather than a reflex mechanism a more likely culprit.

Added in the revised version: Moreover, it has been shown that joint stiffness is reduced during movement compared to holding control (Bennett et al., 1992; Rack and Westbury, 1974). Along similar lines, muscle spindle activity – which may modulate stiffness – scales with extrafusal muscle fiber activity (such as muscle exertion involved in holding) and forces acting through the tendon (Blum et al., 2020). Such observations could, in principle, explain why we were unable to detect a relationship between resting biases and active movement control but we readily found a relationship between resting biases and active holding control: reduced joint stiffness during movement could scale down the influence of resting abnormalities. There are two issues with this explanation, however. First, it is debatable whether this should be considered an alternative explanation per se: stiffness modulation could be, in total or in part, the manifestation of a central movement/posture CST/RST mechanism similar to the one we propose in our conceptual model. For example, (Blum et al., 2020) argue that muscle spindle firing depends on both peripheral and central factors. Second, increased stiffness would not necessarily help detect differences in how active postural control responds to within-resting-posture vs. out-of-resting-posture perturbations. This is because an overall increase in stiffness would likely increase resistance to perturbations in any direction. Interestingly, (Bennett et al., 1992) found that, while stiffness was modulated with elbow angle and gravity during movement, this effect was absent during posture maintenance. This suggests that both the patterns of resting biases we observe in Experiment 1 (tendency towards elbow flexion) and their increase under gravity (i.e. without arm support) cannot be explained by differences in stiffness.

Added in the revised version: An inherent limitation of our finding that resting biases had no clear effect on active movement is that, being a negative finding, it does not prove there is an absence of effect. Nevertheless, our results suggest that any influence of resting biases on active movement may be subtle or suppressed under provided arm support. This contrasts the clear connection between resting biases and active holding control that we also found. Moreover, additional analysis demonstrated that the relative asymmetry in how the motor system responds to perturbations that are against vs. in line with resting biases was greater in the holding compared to the movement case (Figure 8). Another limitation in this comparison is that the perturbations we used to probe active movement vs. active postural control were different from each other, with each perturbation tailored to each modality. Still, our conceptual model does not reject the idea that resting abnormalities may spill into active movement under the right conditions. The exact relationship between these two modalities may be elucidated in further work, potentially by systematically titrating the amount of arm support to find the point at which resting biases begin to substantially affect active movement.

Conclusions

Our examination of the interplay between abnormalities in moving and holding in patients making planar reaching movements suggests the existence of two functional modes of control, likely constructed out of differing configurations of the CST and the RST. The components of the paretic syndrome – loss of dexterity, weakness, abnormal resting posture, and intrusive synergies provide insight into how a normal movement is assembled by descending systems. To the degree that hemiparesis is a consequence of the CST losing, to varying degrees, both its direct control over motor neurons and its modulatory influence on the RST, then physiological and behavioral interventions that augment the residual CST may have a plurality of benefits. In support of this, in a recent study of epidural stimulation of the cervical spinal cord in two chronic stroke patients (the rationale behind the stimulation being to amplify residual CST commands) multiple hemiparetic components –strength, dexterity, synergy abnormalities – improved simultaneously (Powell et al., 2023).